```python
# Initialize Otter
import otter
grader = otter.Notebook("Lab1.ipynb")
```

⚠️**Please first run the code cell above**. This cell initiates the auto-grader within the notebook. Note that the autograder only check if you are giving a **valid** answer, (e.g., type of answer), and thus passing the public tests does not necessarily mean you give the right answer, and consequently get the full points.

# **ESS-314-Lab 1**
* The lab was designed by Ken Creager, Harold Tobin, Anna Ledeczi and Marine Denolle.
* Modified by Barrett Johnson (bnjo@uw.edu); 
* Made autograder available through Otter and GradeScope by Yiyu Ni (niyiyu@uw.edu).
---

## **Lab 1A: Setting up Your Workspace**

This lab focuses on helping students set up their workspace so that they can access the relevant
software packages necessary for future lab work. *You do not need to turn anything in for portion 1A of the lab.*



 ### **Computer Resources**

1. ***UW-JupyterHub*** **(highly recommended)** 

 

UW-IT set up a jupyter notebook for our course with pre-installed packages that you would need during the class. You can start one with the following link in your browser:

https://jupyter.rttl.uw.edu/2023-autumn-ess-314-a/user-redirect/tree

You will need UW ID credentials to get started. You can log in from the ESS Lab computers and from your laptop.

You will be prompted to a Jupyter Notebook in Python. We thrive to use open-source python codes that should be available through our python environment or easy to install using simple commands.



**Now, we will go through an example on how to get the hub ready**.

* Log in to the server
* To add files, you can click the upload button on the top and upload individual files and folders.
* To download lab notebooks, open a **Terminal (command line tool)** window and type 
```
wget  https://github.com/UW-geophysics-edu/ESS314-fall23/raw/main/Lab1.zip
unzip Lab1.zip
```
The lab will be transfered to your environments. 

* The files that you add to the environment will be saved and left as is for future times.

2. ***Google Colab***

Google Colab is a free python notebook environment. It comes in with a built-in python environment. Once you open a notebook in Colab, click ``Connect`` on the top right, select ``Connect to a hosted runtime`` to launch your instance on the free Google Cloud Platform.

To add packages, just use the following command **everytime** you log in (yes, Google don't hold your previous session), for instance:

``!pip install mypackage``

Actually these are **all** packcages during the lab session so save it if you are running notebooks on Google Colab because you would need this in your notebook.
``!pip install otter-grader numpy matplotlib harmonica pyrocko obspy``

⚠️Some packages (yes it's you obspy) may require extra step when running on Google Colab. You might need to restart your run time environment to make it fully functional.



3. ***Your own Local Python***

We created an conda environment in the class Github Page: https://github.com/UW-geophysics-edu/ess314-image/blob/main/environment.yml. Similar to the hub, open a terminal window. Note that you would need to install anaconda/miniconda (python environment & package manager) on your laptop.

Move to the directory of choice and type:
*  ``wget https://raw.githubusercontent.com/UW-geophysics-edu/ess314-image/main/environment.yml``
* Create the conda environment: ``conda env create -f environment.yml``
* start the conda environment: ``conda activate ess314``

Now you are ready to do some pythoning!

### **Getting Familiar with Shell scripting**

Shell scripting allows you to navigate through your computer directories, files, and perform basic manipulation of text files. It can also allows you to run big simulations on large computers. It is extremely flexible!

Start by opening a **Terminal** window. The basic commands you will need are:

* ``pwd`` to see what directory you are in.
* ``cd FOLDER`` to move to another directory
* ``cd ../OTHER_FOLDER`` to move to another directory that is ahead of your current directory. This provides a *relative* path.
* ``cd /home/user/whatever/your/path/is`` is you know the *absolute* path of the directory you want to go to.
* ``ls `` it will list all of the files in the current directory
* ``ls ./*``, exact same as above
* ``ls ../*``, will list the files in the directory ahead of your current directory
* ``mkdir NEW_DIRECTORY`` wil create a new directory
* ``cp file to_new_directory`` will copy ``file`` to another directory
* ``cp -r this_directory to_this_directory``, will copy a folder and everything within it into another directory.


We recommend to take the SoftwareCarpentry classes if you want to learn more. They have sometimes workshop through UW/eScience, or you can follow their lessons on Shell scripting here: 
https://swcarpentry.github.io/shell-novice/

---

## **Lab 1B: Introduction to Python**

Python is an increasingly common scientific tool which we will be utilizing in this course. Lab 1B focuses on gettnig acquainted with Python basics.

**Setting Variables**

 
1.   Many basic commands are available in Python, including arithmetic functions. For example, type 2+2 into a code cell and press SHIFT+ENTER.

```python
2+2
```

2. Arithmetic operators are + for addition, - for subtraction, * for multiplication, and / for division. Raising a value to a power is accomplished by using ** (e.g. 3**2).

3. Type a=2 into the next cell. This will assign the value 2 to the variable a, which will hold this value until it is reassigned. To reassign a, type a=3 in the next cell.

```python
a = 2 
# Make sure to add comments to your code as you go to ensure clarity
```

```python
a = 3 # Here you could remark on what the variable a represents in your code
```

4. Assign a value to b. Now type c = a+b and Python will calculate the value of c. Note that Python will not print out the result of c automatically unless it is told to do so. To see the value of c, type print(c). 

```python
b = 4
c = a + b
c
```

5. There are several kinds of variables in Python. Numbers can be integers, like `a`, whole numbers without a decimal point. Numbers designated as float have a decimal value. For computational purposes, float numbers are more useful. Python also allows the designation of text, or 'strings' as variables. Type `d = 7.2` and `e = 'hello!'` to utilize float and string variables.

```python
d = 7.2
e = 'hello!'
```

6. To see all variables, their type, and their value, use the command %whos. Deleting a single variable can be accomplished using the command del following by the desired variable (e.g., del a). To delete all defined variables, use %reset. Check if the deletions were successful using %whos.

```python
del a 
%whos
```

7. It is important to note that variables in Python can consist of numbers, upper and lowercase letters, and underscores. A variable name cannot start with a number. Variable names are generally more useful when they are more descriptive (e.g., gravity_anomaly rather than ga). 

<!-- BEGIN QUESTION -->

**Question 1a:** What command would be used to clear just the b variable? (1 point)

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 1b:** What command would be used to clear all variables? (1 point)

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 1c:** Which of the following are valid variable names? Define and put your answer in a string variable called `answer_1c`, e.g., `"145"`. (2 points)
1. *_n*
2. *n6*
3. *n_*
4. *6n*
5. *nnnnnnnnnnnnnnnnnnnnn*

```python
# Type your answer here
answer_1c = ...
answer_1c
```

```python
grader.check("q1c")
```

<!-- BEGIN QUESTION -->

**Question 1d:** What error message appears when you try to use an invalid variable name? (1 point)

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**More to read:** 

Actually, there is a whole convention of Python programming called PEP ([Python Enhancement Proposal](https://peps.python.org/pep-0008/)). It defines how you should define variables, constants, functions, and structure of your code. 

There are also tools checking whether your code follows these convention, and helps you correct them. Practice your coding, make the scripts more readible: not only for yourself but also others.

---

**Basic Commands and NumPy**





1.   In order to perform more advanced calculations, certain Python packages must be imported. For example, numerical python is a key package. To import this, use the command 






```python
import numpy as np
```

2. To use numpy, commands must begin with np. For example, type np.pi and you will see the result printed to the screen. To round this number to a certain number of decimal places, assign the result to a variable, then use the function round, for example:


```python
pi = np.pi
round(pi, 2)
```

3. In numpy, log(x) calculates the natural log of a number. log10(x) will return the log base 10 and log2(x) will return the log base 2. 


```python
np.log10(10)
```


4. The command np.sqrt() can be used to calculate the square root of positive numbers.


```python
np.sqrt(3)
```


5. Trigonometric functions in NumPy take angles in radians, so make sure your input angles are not in degrees. You can also use the built-in NumPy conversion between degrees and radians.

```python
print(np.sin(np.pi/2))
print(np.sin(np.deg2rad(90))) # You can also use np.rad2deg() for other calculations
```

6. Inverse trigonometric functions can be calculated using np.arcsin(), np.arccos(), and np.arctan().

```python
print(np.arcsin(np.sin(np.pi/3)))
print(np.pi/3)
```

<!-- BEGIN QUESTION -->

**Question 2a:**  What is the command to return the value of e (Euler’s number)? (1 point)

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 2b:** What is the sin of 270 degrees? (1 point)

```python
# Type your answer here
answer_2b = ...
answer_2b
```

```python
grader.check("q2b")
```

**Question 2c:** What is arccos(0.5) in degrees? (1 point)

```python
# Type your answer here
answer_2c = ...
answer_2c
```

```python
grader.check("q2c")
```

**Question 2d:** What is arccos(0.5) in radians? (1 point)

```python
# Type your answer here
answer_2d = ...
answer_2d
```

```python
grader.check("q2d")
```

**Vectors and Matrices**

NumPy can be utilized to create "arrays", which function as vectors or matrices. The size of a matrix with r rows and c columns is 𝑟 × 𝑐. Think of a row vector as a matrix with 𝑟 = 1, and a column vector as a matrix with 𝑐 = 1.

1.  Use the command np.array() to create row and column vectors as follows:

```python
row_vector = np.array([[1, 2, 3]])
print(row_vector)

col_vector = np.array([[1], [2], [3]])
print(col_vector)

# You can check that these vectors are indeed the dimensions you seek by using the following command:

print(row_vector.shape) # Should be (1, 3) for (r, c)
print(col_vector.shape) # Should be (3, 1) for (r, c)
```

2. Selecting specific values in a vector or matrix is known as indexing. Start by creating a 3 x 4 matrix as follows.

```python
matrix = np.array([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12]])
print(matrix)
```

To select elements within the matrix, specify their location within the rows and columns, making sure to begin counting at 0. If you are used to MATLAB, it is important to remember that **Python indexing starts at 0, not 1**. Here are some examples of indexing:

```python
print(matrix[1,2]) # The number 7
print(matrix[2,0]) # The number 9
print(matrix[0, :]) # The first row
column1 = matrix[:, 0] # The first column - but note it only has one dimension
print(column1)
column1.shape # but note it only has one dimension and is not technically a 'row' or 'column'

matrixT = np.transpose(matrix) # Transpose a matrix using this command
print(matrixT)
```

3. There are other ways to create more specialized kinds of arrays, such as those made out of zeroes, ones, or a linearly spaced array.

```python
print(np.zeros([3,1]))

print(np.ones([4,1]))

print(np.linspace(0,10,5)) # Specify start, stop, and the number of elements for a linearly spaced array

print(np.arange(0, 10, 2)) # Specify start, stop, and spacing

# Note the difference between the following: 

print(np.linspace(0,10,10))
print(np.linspace(0,10,11))
```

4. You can concatenate or truncate vectors once they are created. The two arrays must have the same shape, except in the direction of the axis. 

```python
g = np.ones([4,3])
print(g)

h = np.zeros([2,3])
print(h)

i = np.concatenate((g,h)) # axis is assumed 0 unless specified 
print(i)


j = np.ones([1,2])
print(j)

k = np.zeros([1, 3])
print(k)

l = np.concatenate((j, k), axis=1)
print(l)

m = np.concatenate((k, j), axis=1) # reverses the order from above
print(m)
```

```python
# Use these for Question 3

v = np.array([[1, 2, 3]])
w = np.array([[1], [2], [3]])
```

**Question 3a:** What is the difference between v and w? (1 point)

```python
# Type your answer here
answer_3a = ...
answer_3a
```

```python
grader.check("q3a")
```

<!-- BEGIN QUESTION -->

**Question 3b:** Explain the function of np.transpose(). (1 point)

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 3c:** Create a matrix with three columns, each of which is a copy of vector w. (1 point)

```python
# Type your answer here
answer_3c = ...
answer_3c
```

```python
grader.check("q3c")
```

<!-- BEGIN QUESTION -->

**Question 3d:** How to select the entire third column of the matrix you just created? (1 point)

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Operations with Vectors and Matrices**



```python
v = np.array([[1, 2, 3]])
w = np.array([[1], [2], [3]])
p = ([[3, 4, 5, 6]])
M = np.array([[1, 2, 3], [4, 5, 6], [7, 8, 9]])
N = np.concatenate((v, v, v))
```

1. Adding a scalar to a vector or matrix is the same as adding two scalar values. Similarly,
multiplying a vector or matrix by a scalar value looks just like multiplying two scalars.
a. To add 2 to all elements of your vector v, type v+2.
b. To multiply each element of matrix M by 3, type M*3.

```python
print(v+2)
print(M*3)
```

2. Adding two matrices may work differently than you expect. Try adding v+N, v+p, v+w

```python
v+N
```

3. To perform matrix multiplication, just multiply the two matrices by each other. Recall
that to multiply two matrices together, the number of columns in the first must equal
the number of rows in the second.
a. Try multiplying vector v by vector w. Make sure and show that they have the complementary dimensions.

```python
print(v.shape)
print(w.shape)
v*w
```

4. For square matrices, you can also raise the matrix to a power using the same function as
you did with scalar values. Square the matrix M using the command M**2, then review
the output.

```python
M**2
```

5. Operations that apply to each element of the matrix individually are sometimes called array operations or element-wise operations. Many functions, such as trig functions, will operate on a matrix element-by-element. For example, try np.sin(N).

```python
np.sin(N)
```

<!-- BEGIN QUESTION -->

**Question 4:** What does Python do when you ask it to add v+M? What about p+M? Describe as best you can from your own experiments the logic that Python uses to add two matrices together. Include any relevant errors that might occur. (1 point)

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

**Plotting**



1.   Visualization is important! Plotting in Python requires importing another package, known as matplotlib. You can find tutorials and self-teach yourself about plotting in Python on the Matplotlib tutorial page: https://matplotlib.org/stable/tutorials/index.html



```python
import matplotlib.pyplot as plt 
# You will have to run this import cell each time you open the notebook
```

2. Create a row vector x with the numbers 1 through 50. Then, use the sine function to vector y where each element is the sine of the corresponding value in x.

```python
x = np.linspace(1,50,50)
y = np.sin(x)
```

3. Create a simple plot using matplotlib.

```python
plt.figure(figsize=(10, 5), dpi=100) # define the figure size and resolution
plt.plot(x,y,'g:') # 'g:' defines the color (g) and the pattern (:) of the lines here. 
# You can search for color legends and more line patterns. - is a line, o is dots. 
plt.xlabel('x label')
plt.ylabel('y label')
plt.title("Example Plot 1")
plt.show()
```

4. Create a figure with 2 subplots.

```python
z = np.cos(x)
```

```python
plt.figure(figsize=(10, 5), dpi=100) # control how many rows and columns here

plt.subplot(1, 2, 1) # split figure into 1 row, 2 columns, and plot on the 1st subplot
plt.plot(x,y, 'r-', label='sin(x)')
plt.title('Example Plot 2a')
plt.xlabel('This is x (radians)')
plt.ylabel('Sine of x')
plt.legend()

plt.subplot(1, 2, 2) # split figure into 1 row, 2 columns, and plot on the 2nd subplot
plt.plot(x, z, 'bo', label='cos(x)')
plt.title('Example Plot 2b')
plt.xlabel('This is x (radians)')
plt.ylabel('Cosine of x')
plt.legend()

plt.tight_layout() # This will add space between your subplots (try commenting out this line!)
plt.show()

# Use the following command to save the image
plt.savefig('example.png')
```

5. Redefine x with twice as many points (i.e., define x as a row vector from 1 to 50, but include points at every half integer). Try to plot x and y again. What's wrong?

```python
x = np.linspace(1,50,99)
fig = plt.plot(x, y, 'b-')
```

<!-- BEGIN QUESTION -->

**Question 5a:** Redefine y as appropriate. Keep iterating on the resolution of x and y until you're satisfied that the graph looks like a sine wave. Add a title and axis labels. (1 point)

Hint: you can use ?np.linspace to check manual of the function quickly.

```python
?np.linspace
```

```python
# Type your answer here
x = ...
y = ...
plt.figure(figsize=(10, 5), dpi=100)
plt.scatter()
...
```

<!-- END QUESTION -->

**Question 5b:** Create a row vector from 0 – 100, counting by 5 (i.e. [0 5 10 15 20 … 100])? Define and put your answer in a string variable called `x`. (1 point)

```python
# Type your answer here
x = ....
x
```

```python
grader.check("q5b")
```

<!-- BEGIN QUESTION -->

**Question 5c:** Plot x and y with maroon upward-facing triangles? (1 point)

Hint: check the manual here: 
- plt.plot: https://matplotlib.org/stable/api/_as_gen/matplotlib.pyplot.plot.html
- Markers: https://matplotlib.org/stable/api/markers_api.html#module-matplotlib.markers

```python
# Type your answer here
plt.figure(figsize = (10, 4), dpi = 100)
...
```

<!-- END QUESTION -->

### Plot two functions on the same plot. 

**Question 6a:** Define a new row vector `x` where the values range from 1 to 90 **in degree** with a 0.5 interval (1, 1.5, 2, ...), and **convert the vector to radian**. (1 point)

```python
# Type your answer here
x = ...
x
```

```python
grader.check("q6a")
```

**Question 6b:** Define a vector y that is equal to x^2 + 50cos(x + 30). (1 point)

```python
# Type your answer here
y = ...
y
```

```python
grader.check("q6b")
```

**Question 6c:** Define a vector z that is equal to 1.5(tan(y)). (1 point)

```python
# Type your answer here
z = ...
z
```

```python
grader.check("q6c")
```

<!-- BEGIN QUESTION -->

**Question 6d:** Plot y vs. x and z vs. x on the same graph. (2 points)
* Plot the y function as a blue solid line, and the z function as a red dotted line.
* Add a title, and label the x axes `x`, `y(x)` and `z(y(x))`, respectively.

```python
# Type your answer here
plt.figure(figsize = (10, 4), dpi = 100)
plt.plot(...)
plt.legend()
plt.xlabel(...)
...
```

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are the lab 1 of ESS 314 Autumn 2023.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1c (points: 2)
```
>>> type(answer_1c) == str
True
```


Tests for q2b (points: 1)
```
>>> np.issubdtype(type(answer_2b), np.number)
True
```


Tests for q2c (points: 1)
```
>>> np.issubdtype(type(answer_2c), np.number)
True
```


Tests for q2d (points: 1)
```
>>> np.issubdtype(type(answer_2d), np.number)
True
```


Tests for q3a (points: 1)
```
>>> type(answer_3a) == np.ndarray
True
```


Tests for q3c (points: 1)
```
>>> type(answer_3c) == np.ndarray
True
```


Tests for q5b (points: 1)
```
>>> type(x) == np.ndarray
True
```


Tests for q6a (points: 1)
```
>>> type(x) == np.ndarray
True
```


Tests for q6b (points: 1)
```
>>> type(y) == np.ndarray
True
```


Tests for q6c (points: None)
```
>>> type(z) == np.ndarray
True
```
