```python
# Initialize Otter
import otter
grader = otter.Notebook("Lab4.ipynb")
```

# **ESS 314: Lab 4**

**Plotting travel-time curves from reflected waves in Python**

Imagine a source 10 m from the first receiver (i.e., geophone) in a 10-geophone line.
The spacing between geophones is 10 m. We also know that the thickness of the layer
closest to the surface is 20 m. The velocity of layer 1 is 2000 m/s and the velocity of
layer 2 is 4000 m/s.

![](source_receiver_geometry.png)

### Import packages that are used in the notebook
You would need to do this everytime you start a new notebook.

```python
import numpy as np
import matplotlib.pyplot as plt
from pyrocko import cake
import pandas as pd
```

We use the same code from lab 3 that defines a layered model, with one 20 meters homogeneous layer. The Vp of layer1 is `2000 (m/s)`. **We still won't use S-wave phases during this lab.**

```python
m_layer1 = cake.Material(vp=2000.)
m_layer2 = cake.Material(vp=4000.)

model = cake.LayeredModel()
model.append(cake.Surface(z=0., mbelow=m_layer1))
model.append(cake.HomogeneousLayer(0., 20., m_layer1, name='layer1'))
model.append(cake.Interface(z=20., mabove=m_layer1, mbelow=m_layer2, name="interface1"))
```

Similar to lab 3, we define the source to be sightly below the surface, and receivers arranges at 10, 20, 30, ... 100 meters.

```python
source_depth = 0.01 # in meters
distances = np.arange(10, 101, 10) * cake.m2d  # convert meter to degree
```

Here we define the phase that we calculate the travel time. There is a subtle difference in the phase defined. Remember that in lab 3 we define the refraction wave along interface 1 as `Pv_(interface1)p`. Here, for the reflection wave, there is no `_` (underscore), indicating that the phase bounds back immediately after touching the interface, making it the reflection.

```python
phase = cake.PhaseDef('Pv(interface1)p')
```

```python
# save the travel-time
x_reflect_cake = [];
t_reflect_cake = [];
for arrival in model.arrivals(distances, phases=phase, zstart=source_depth):
    print('%13g %13g' % (arrival.x*cake.d2m, arrival.t))
    x_reflect_cake.append(arrival.x*cake.d2m)
    t_reflect_cake.append(arrival.t)
```

<!-- BEGIN QUESTION -->

📍**Question 1a:** Notice that this is the same subsurface configuration as in the refraction lab. Add the **time-distance curve for the reflected wave to your plot from Lab 3, question 1b** – this should only require a line or two added to your previous code. Make sure to label axes and add a legend. (2 points)

```python
# Type your answer here, replace "..." with your code
plt.figure(figsize=(10, 5), dpi=100)
plt.plot(..., ...)
...
plt.show()
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

📍**Question 1b:** Make the geophone spacing 20 m instead of 10 m, copy your code, and rerun. (1 point)

```python
# Type your answer here, replace "..." with your code
plt.figure(figsize=(10, 5), dpi=100)
plt.plot(..., ...)
...
plt.show()
```

<!-- END QUESTION -->

Notice that as distance from the shot increases, the arrival times of the direct wave and the reflected wave get closer together. Think about this in terms of a triangle formed by the half of the surface distance from the shot to the geophone (x), the thickness of the layer (h), and the ray path to the boundary layer (r). The layer thickness is constant but increasing the value of x also changes the length of r and the angle of incidence.

📍**Question 1c:** Explain why the arrival times of the reflected wave approaches that of the direct wave using the diagram above as a template. Clearly reference the ray paths of both direct and reflected waves from the diagram above in your explanation. (1 point)

_Type your answer here, replacing this text._

📍**Question 1d:** Explain why travel times of the critically refracted wave are similar to the reflected wave on geophones closer to the shot, but arrive earlier than reflected waves for geophones further from the shot. Again, explain this with clear references to your two plots in a paragraph similar to the one above. Think about the relative distances each ray travels in each material. (1 point)

_Type your answer here, replacing this text._

---

### **Use pyrocko.cake to calculate time-distance curves**

📍**Question 2a:** We create a new 3-layer model, each layer is 50 meters thick. Let V1 = 1000 m/s, V2 = 2000 m/s, and V3 = 4000 m/s. Use 50 geophones with 10 meter spacing. Fill in the cake model below and calcuate the travel time. Plot the reflection wave from each interface. Make sure your plot looks like the one below. (2 points)
 
 ![](plot_arrivals.png)

```python
# Type your answer here, replace "..." with your code
m_layer1 = cake.Material(vp = ...)
m_layer2 = cake.Material(vp = ...)
m_layer3 = cake.Material(vp = ...)
m_layer4 = cake.Material(vp = ...) # does this matter?

model = cake.LayeredModel()
model.append(cake.Surface(z=0., mbelow=m_layer1))
model.append(cake.HomogeneousLayer(0., ..., m_layer1, name='layer1'))
model.append(cake.Interface(z=..., mabove=m_layer1, mbelow=m_layer2, name="interface1"))
model.append(cake.HomogeneousLayer(..., ..., m_layer2, name='layer2'))
model.append(cake.Interface(z=..., mabove=m_layer2, mbelow=m_layer3, name="interface2"))
model.append(cake.HomogeneousLayer(..., ..., m_layer3, name='layer3'))
model.append(cake.Interface(z=..., mabove=m_layer3, mbelow=m_layer4, name="interface3"))
print(model)
```

```python
source_depth = 0.01 # in meters
distances = np.arange(10, 501, 10) * cake.m2d  # convert meter to degree
```

We calculate the travel time of reflection wave of each interface, separately.

```python
####### reflection of interface 1
phase = cake.PhaseDef('Pv(interface1)p')
x1_cake = [];
t1_reflect_cake = [];
for arrival in model.arrivals(distances, phases=phase, zstart=source_depth):
    x1_cake.append(arrival.x*cake.d2m)
    t1_reflect_cake.append(arrival.t)

####### reflection of interface 2
phase = cake.PhaseDef('Pv(interface2)p')
x2_cake = [];
t2_reflect_cake = [];
for arrival in model.arrivals(distances, phases=phase, zstart=source_depth):
    x2_cake.append(arrival.x*cake.d2m)
    t2_reflect_cake.append(arrival.t)

####### reflection of interface 3
phase = cake.PhaseDef('Pv(interface3)p')
x3_cake = [];
t3_reflect_cake = [];
for arrival in model.arrivals(distances, phases=phase, zstart=source_depth):
    x3_cake.append(arrival.x*cake.d2m)
    t3_reflect_cake.append(arrival.t)
```

```python
plt.figure(figsize=(10, 6), dpi = 100)
plt.plot(t1_reflect_cake, x1_cake, linewidth=2)
plt.plot(t2_reflect_cake, x2_cake, linewidth=2)
plt.plot(t3_reflect_cake, x3_cake, linewidth=2)
plt.gca().invert_yaxis()
plt.grid(True)
plt.xlabel("Time (s)", fontsize = 14)
plt.ylabel("Distance (m)", fontsize = 14)
```

<!-- BEGIN QUESTION -->

📍**Question 2b:** Answer the questions below.

1. For the plot above, indicate the layer off which each reflected wave bounces (blue line, orange line, and green line). (1 point)

2. On geophones far away from the shot, why does the second reflected wave arrive before the wave reflected off the first layer? (1 point)

**For the questions below, remember that if you change the thickness of the first layer, you need to change the depth of layers and interfaces below.**

3. If we decrease the thickness of layer 1 to 20 m, how does this affect the arrival times of the first reflected wave? Why? (2 points)

4. Change the thickness of the first layer back to 50 m and decrease the thickness of the middle layer to 10 m. How does this affect the arrivals of the second reflected wave? Why?  (2 points)

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

📍**Question 3:** The basic reflection travel-time equation for a single layer is similar to the equation of a hyperbola:

$ 𝑡 = \dfrac{\sqrt{x^2+4h_1^2}}{V_1}$ 

To simplify analysis, square both sides and write it in the form of an equation for a straight line. This is called the $x^2-t^2$ method.

Rewrite the reflection travel-time equation in linear form $y=mx +b$ by squaring both sides, expressing $𝑡^2$ as a linear function of $𝑥^2$. Don't worry about rendering it like the equation above. Show your work. (2 points)

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

📍**Question 4:**

Calculating the travel time of waves reflected at deeper layers is complicated because those waves are refracted at shallower boundaries. One way to deal with this complication is Green’s Method, which assumes straight-line paths through all boundaries and solves for velocities and thicknesses. This method is illustrated in the figure below.


![](funny_rays.png)
1. Draw the actual ray paths of the reflected waves on the diagram for each ray path. You can draw this on a piece of paper or by editing a screenshot and submit this to GradeScope with your assignment named `fig_4.png`. (1 point)

2. Use your plot to explain below why Green’s Method may not be the most accurate approximation. (1 point)

_Type your answer here, replacing this text._

<!-- END QUESTION -->

Another way to solve velocities and thicknesses is to use the Dix equation, which uses
the root-mean-square velocity ($V_{rms}$). This was covered in lecture.

---

### **Modeling the Cascadia subduction zone**

The Cascadia Subduction Zone (CSZ), where the Juan De Fuca plate dives beneath the North American plate, runs from northern California to Vancouver Island in Canada. This subduction zone is capable of generating large (>M9.0) earthquakes. Since these earthquakes pose a serious threat to cities and towns along the West Coast, we want to learn as much as we can about the CSZ. Shallower earthquakes can cause much more damage than a deep earthquake of the same  size, so determining the depth of the plate interface, where megathrust earthquakes occur, is important. 

In an attempt to figure out the depth of the plate interface, a set of 50 geophones spaced 1000 meters apart was set up on the Olympic Peninsula in an east-west trending line. A seismic reflection study recorded P-wave arrivals; the data is in the `SubductionData.txt` file.

Here, we load the data using the package called pandas (short name as `pd`).

```python
df = pd.read_csv("./SubductionData.txt", skiprows=3, names=['RecID', 'x', 't1', 't2', 't3'])
df.head() # this function shows the first several rows of the table
```

Still, we define the source and receivers location. 

```python
source_depth = 0.01 # in meters
distances = np.linspace(0, 49000) * cake.m2d  # convert meter to degree
```

And we plot the data from this experiment. Let's make it look nice with scatter points. We use `plt.scatter` here since we only have limited and sparse observation from the experiment.

```python
plt.figure(figsize=(10, 6), dpi=100)
plt.scatter(df['x'], df['t1'], s=40, label="t1", marker='x', zorder=10)
plt.scatter(df['x'], df['t2'], s=40, label="t2", marker='x', zorder=10)
plt.scatter(df['x'], df['t3'], s=40, label="t3", marker='x', zorder=10)
plt.grid(True)
plt.legend(fontsize=14)
plt.xlabel("Distance (m)", fontsize=12)
plt.ylabel("Time (s)", fontsize=12)
```

<!-- BEGIN QUESTION -->

📍**Question 5:**
How many layers (or interfaces) does this reflection study show, judging by the number of reflections? Put your answer as an integer in the variable below. (1 point)

```python
# Type your answer here, replace "..." with your answer, e.g., 3
answer_5 = ...
answer_5
```

```python
grader.check("q5")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

📍**Question 6:**

1. To analyze these data, we will need to plot the data on an $x^2-t^2$ plot. Modify the code below to show data in $x^2-t^2$ space. (1 point)

2. Define a cake model to calculate the travel time of reflection wave off the first interface. It is OK to start with an one-layer model. But note that we don't know the P-wave velocity of the layer, so we need to guess. Here I give you an inital guess of `Vp = 3000 m/s` and `15 km` depth. After you have the travel time of this reflection wave (from your guessed model), add it to the plot below using dashed line `--`. (1 point)

3. Does the line fit the observation in terms of the intercept and the slope? What does this suggest about our initial guess of the model? (1 point)

4. Based on the equation from Question 3, if you move the line for layer 1 up without changing the slope, do you expect the thickness or velocity of your model to change? Why? (1 point)

5. Based on the equation from Question 3, describe how you think changing the slope of the line for layer 1 will affect the thickness or velocity of the layer. (1 point)

6. Based on your answer to question c, d, and e, change the parameter of layer 1 (velocity and depth) until it fits the data. It does not need to be perfectly match. (1 point)

```python
m_layer1 = cake.Material(vp = 3000.)     # layer 1 P-wave velocity
m_layer2 = cake.Material(vp = 9999.)     # layer 2 P-wave velocity
m_layer3 = cake.Material(vp = 9999.)     # layer 3 P-wave velocity
m_layer4 = cake.Material(vp = 9999.)     # layer 4 P-wave velocity

d_interface1 = 15000.                    # interface 1 depth in meter
d_interface2 = 99999.                    # interface 2 depth in meter
d_interface3 = 99999.                    # interface 3 depth in meter

########## Change the parameter above ##########

########## You don't need to chagne the code below ##########
model = cake.LayeredModel()
model.append(cake.Surface(z = 0., mbelow = m_layer1))
model.append(cake.HomogeneousLayer(0., d_interface1, m_layer1, name='layer1'))
model.append(cake.Interface(z = d_interface1, mabove = m_layer1, mbelow = m_layer2, name = "interface1"))
model.append(cake.HomogeneousLayer(d_interface1, d_interface2, m_layer2, name='layer2'))
model.append(cake.Interface(z = d_interface2, mabove = m_layer2, mbelow = m_layer3, name = "interface2"))
model.append(cake.HomogeneousLayer(d_interface2, d_interface3, m_layer3, name='layer3'))
model.append(cake.Interface(z = d_interface3, mabove = m_layer3, mbelow = m_layer4, name = "interface3"))

phase = cake.PhaseDef('Pv(interface1)p')
x1_cake = [];
t1_reflect_cake = [];
for arrival in model.arrivals(distances, phases=phase, zstart=source_depth):
    x1_cake.append(arrival.x*cake.d2m)
    t1_reflect_cake.append(arrival.t)
x1_cake = np.array(x1_cake)
t1_reflect_cake = np.array(t1_reflect_cake)

phase = cake.PhaseDef('Pv(interface2)p')
x2_cake = [];
t2_reflect_cake = [];
for arrival in model.arrivals(distances, phases=phase, zstart=source_depth):
    x2_cake.append(arrival.x*cake.d2m)
    t2_reflect_cake.append(arrival.t)
x2_cake = np.array(x2_cake)
t2_reflect_cake = np.array(t2_reflect_cake)

phase = cake.PhaseDef('Pv(interface3)p')
x3_cake = [];
t3_reflect_cake = [];
for arrival in model.arrivals(distances, phases=phase, zstart=source_depth):
    x3_cake.append(arrival.x*cake.d2m)
    t3_reflect_cake.append(arrival.t)
x3_cake = np.array(x3_cake)
t3_reflect_cake = np.array(t3_reflect_cake)
```

```python
# Type your answer here, replace "..." with your code
plt.figure(figsize=(10, 6), dpi=100)
plt.scatter(..., ..., s=40, label="t1", marker='x', zorder=10)
plt.scatter(..., ..., s=40, label="t2", marker='x', zorder=10)
plt.scatter(..., ..., s=40, label="t3", marker='x', zorder=10)

plt.plot(..., ..., '--', label='t1_cake')
plt.plot(..., ..., '--', label='t2_cake')
plt.plot(..., ..., '--', label='t3_cake')

plt.grid(True)
plt.legend(fontsize=14, loc = 'upper left')
plt.xlabel("Distance^2 (m^2)", fontsize=12)
plt.ylabel("Time^2 (s^2)", fontsize=12)
plt.show()
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

📍**Question 7:** Keep working on the plot above. 

1. Once you are happy with the paramemter of the first layer (depth and vp), let's work on fitting the second layer. Add the plot from second layer reflection as dashed line. (1 point)

2. Change the parameter of the second layer in various ways. Note that changing the line for the first layer changes the whole model, but that is not the case for the second layer. Explain why this is the case. (1 point)

3. Change the parameter of the second layer reflection until it match the data. (1 point)

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

📍**Question 8:** Keep working on the plot above. Once you are happy with the paramemter of the first and the second layer (depth and vp), let's work on fitting the third layer. Add the plot from third layer reflection as dashed line. Change the parameter of the third layer reflection until it match the data. (2 points)

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

📍**Question 9:** 

1. Based on your final model, discuss the section below, indicating where you think the plate interface is. (1 point)

2. Write a paragraph about the validity of this section, using the points below to develop your discussion and referencing your diagram and its labels. (2 points)

    i. Keep in mind that both the upper plate and the lower plate both have a crust and subsequent mantle beneath, so think about which layer is which using the typical mantle P-wave velocity around 8 km/s.

    ii. What, if anything, can you say about the velocity of the deepest layer based on reflection data alone?

    iii. Use the table below to determine what types of rocks make up the crusts of the upper and lower plates. Make sure your interpretation makes sense given that crust of the lower plate is oceanic and the crust of the upper plate is continental.

| Rock Type      | P wave velocity (m/s) |
| ---------------| ----------- |
| Sandstone and shale | 2000-4500  |
| Limestone | 2000-6000        |
| Granite   | 5000-6000     |
| Basalt    | 5400-6400      |

_Type your answer here, replacing this text._

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are the lab 4 of ESS 314 Autumn 2023.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q5 (points: 1)
```
>>> np.issubdtype(type(answer_5), np.number)
True
```
