```python
# Initialize Otter
import otter
grader = otter.Notebook("Lab6.ipynb")
```

# **ESS 314: LAB 6**
## **Gravity**

**Let's first import the python libraries that are required during the lab.**

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt

# new packages for gravity & magnetic modeling
import harmonica as hm
import verde as vd
```

<!-- BEGIN QUESTION -->

📍**Question 1:** What does the Bouguer correction correct for? (1 point)

_Type your answer here, replacing this text._

<!-- END QUESTION -->

📍**Question 2:** The density of the surrounding rock is 2600 $kg/m^3$ and the density of a body of interest is 2800 $kg/m^3$. What is the density contrast in $kg/m^3$? (Make sure you get the sign right.)

```python
# Type your answer here, replacing "..." with your answer
contrast = ...
contrast
```

```python
grader.check("q2")
```

---
#### **Learning to use Harmonica**

**Harmonica** is a Python library for processing and modeling gravity and magnetic data. It includes common processing steps, like calculation of Bouguer and terrain corrections, reduction to the pole, upward continuation, equivalent sources, and more. There are forward modeling functions for basic geometric shapes. You can find more documentation about Harmonica here: https://www.fatiando.org/harmonica/v0.6.0/

To use Harmonica, we first need to define the coordinate system for observation (3-dimentional coordinate system). Scripts below define a `region` of 100x100 meters: 100 meters in both x and y direction (horizontal). The `shape` parameter defines who many points to measure in the defined region, which in this case, 101 measurement in both x and y direction. Then we would have gravity measurement at `x = [0,1,2,...,100]` and `y = [0,1,2,...,100]`. We also define the height where the measurement should be made, which is important for gravity modeling.

```python
region = (0, 100, 0, 100) # (x_min, x_max, y_min, y_max), all in meters
shape = (101, 101)        # (n_x, n_y)
height = 0                # in meters
coordinates = vd.grid_coordinates(region, shape=shape, extra_coords=height)
```

Then, we define the density model that gives gravity response. Here, we define a 3-dimentional block with some density contrast to the background. **Harmonica assumes zero density of the background material, so when defining the body, only use the density constrast.** For example here, we define a block that is 20 meters width in x (left edge at 40 $m$, right edge at 60 $m$), 20 meters width in y (left edge at 40 $m$, right edge at 60 $m$), and 10 meters thick (upper edge at -10 $m$, lower edge at -20 $m$). The density contrast is 200 $kg/m^3$.

```python
block = [40, 60, 40, 60, -20, -10] # (x_min, x_max, y_min, y_max, z_min, z_max), all in meters
density_contrast = 200.            # unit in kg/m^3
```

Using the paramter above, we can modeling the gravity field anomaly ($g_z$) from the higher density block.

```python
result_q3 = hm.prism_gravity(coordinates, block, density_contrast, field="g_z")
print(f"the shape of result is {result_q3.shape}")
```

We can plot the gravity anomaly map (2D). We can also plot the anomaly profile (1D, along the red/yellow/green dashed line in the map).

```python
plt.figure(figsize=(14, 5), dpi=100)
plt.subplot(1,2,1)
plt.title("gravity anomaly map")
plt.pcolormesh(coordinates[0], coordinates[1], result_q3)
plt.hlines(50, 1, 100, color='r', linestyle='--')
plt.hlines(30, 1, 100, color='y', linestyle='--')
plt.hlines(10, 1, 100, color='g', linestyle='--')
plt.colorbar(label="mGal")
plt.xlabel("x")
plt.ylabel("y")

plt.subplot(1,2,2)
plt.title("gravity anomaly profile")
plt.plot(result_q3[50], color='k', linestyle='-')
plt.plot(result_q3[30], color='k', linestyle=':')
plt.plot(result_q3[10], color='k', linestyle='--')
plt.xlabel("x")
```

<!-- BEGIN QUESTION -->

📍**Question 3:**

a. Describe the shape of the curve in the gravity anomaly profile. (1 point)

b. The solid line on the profile plot corresponds to which profile on the map (red, yellow, or green line)? How about the dotted line and the dashed line? (2 points)

c. Where is the peak of the curve relative to the block? (1 point)

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

📍**Question 4:** Give your answers based on the anomaly map. You can use `np.max()` to fine the maximum value in the `result_q3`. (2 points)

a. What is the peak anomaly when the density contrast is 200 $kg/m^3$?

b. If you change the contrast density to 100 $kg/m^3$, what will be the peak anomaly?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

📍**Question 5:**

a. What happens to the curve if you change the contrast density contrast to negative? (1 point)

b. If the density contrast is negative, is the body more or less dense than the surrounding host rock? (1 point)

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

📍**Question 6a:** Remember that we modelled the gravity anomaly for a block with density contrast of $200 kg/m^3$. We now increase the density contrast to 400 $kg/m^3$. Approximately how deep does **the top of the body** (z_max) have to be for the peak anomaly to be the same as it was when the density contrast was 200 $kg/m^3$? You can modify the depth of the block below, but please keep the thickness of the body at a constant 10 $m$. (2 points)

```python
region = (0, 100, 0, 100)
shape = (101, 101)
height = 0
coordinates = vd.grid_coordinates(region, shape=shape, extra_coords=height)
```

```python
# Type your answer here, replacing "..." with your answer
block =      [40, 60, 40, 60, ..., ...]   # (x_min, x_max, y_min, y_max, z_min, z_max), all in meters
density_contrast = ...                    # unit in kg/m^3

result_q6a = hm.prism_gravity(coordinates, block, density_contrast, field="g_z")
print(f"the maximum gravity anomaly is {np.max(result_q6a)} mGal")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

📍**Question 6b:** We also reduce the density contrast to 150 $kg/m^3$. Approximately how shallow does **the top of the body** (z_max) have to be for the peak anomaly to be the same as it was when the density contrast was 200 $kg/m^3$? You can modify the depth of the block below, but please keep the thickness of the body at a constant 10 $m$. (2 points)

```python
# Type your answer here, replacing "..." with your answer
block =      [40, 60, 40, 60, ..., ...]   # (x_min, x_max, y_min, y_max, z_min, z_max), all in meters
density_contrast = ...                    # unit in kg/m^3

result_q6b = hm.prism_gravity(coordinates, block, density_contrast, field="g_z")
print(f"the maximum gravity anomaly is {np.max(result_q6b)} mGal")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

📍**Question 7:** 

We plot the anomaly profile from `result_q3` (corresponds to contrast 200 $kg/m^3$), together with `result_q6a` and `result_q6b` (corresponds to density contrast 400 and 150 $kg/m^3$, respectively). Use the figure below to explain how can you distinguish a shallow, dense body from a deeper, denser body. (2 points)

_Type your answer here, replacing this text._

```python
plt.figure(figsize=(10,6))
plt.plot(result_q3[50, :], c='k')
plt.plot(result_q6a[50, :], c='k')
plt.plot(result_q6b[50, :], c='k')
plt.grid(True)
plt.xlabel("y (m)", fontsize=14)
plt.ylabel("gravity anomaly (mGal)", fontsize=14)
```

<!-- END QUESTION -->

---

<!-- BEGIN QUESTION -->

📍**Question 8:** 

Based on your answer to question 6, you should have noticed that the results from gravity surveys are non-unique. Multiple subsurface configurations can produce very similar anomaly curves.

Load the file `GravityData.txt`. Create two different subsurface models to fit the observations (i.e., add bodies to the subsurface so that the calculated gravities fit the observations). You can use different density contrasts, shapes, and even add more than 1 body. There are infinite answers. (2 points for each modeling)

```python
# Read the data
df = pd.read_csv("./GravityData.txt", names=["measurement", "x", "g_z"])
df.head()        # show the first 5 rows
```

We use the same coordinate setting as above.

```python
region = (0, 100, 0, 100)
shape = (101, 101)
height = 0
coordinates = vd.grid_coordinates(region, shape=shape, extra_coords=height)
```

_Type your answer here, replacing this text._

```python
# Type your answer here, replacing "..." with your answer
blocks = [..., ..., 0, 100, ..., ...]   # (x_min, x_max, y_min, y_max, z_min, z_max), all in meters
density = ...                           # density contrast 

result_q9a = hm.prism_gravity(coordinates, blocks, density, field='g_z')
print(f"the maximum gravity anomaly is {np.max(result_q9a)} mGal")
```

```python
# Type your answer here, replacing "..." with your answer
blocks = [[..., ..., 0, 100, ..., ...], # (x_min, x_max, y_min, y_max, z_min, z_max) of BLOCK 1
          [..., ..., 0, 100, ..., ...]] # (x_min, x_max, y_min, y_max, z_min, z_max) of BLOCK 2
density = [...,                         # density contrast for BLOCK 1
           ...]                         # density contrast for BLOCK 2

result_q9b = hm.prism_gravity(coordinates, blocks, density, field='g_z')
print(f"the maximum gravity anomaly is {np.max(result_q9b)} mGal")
```

```python
plt.figure(figsize=(10, 5))
plt.plot(df['x'], df['g_z'], label = 'data', linewidth = 2)
plt.plot(df['x'], result_q9a[50, :], label = 'my model 1')
plt.plot(df['x'], result_q9b[50, :], label = 'my model 2')
plt.xlabel("x (m)", fontsize=14)
plt.ylabel("gravity anomaly (mGal)", fontsize=14)
plt.grid(True)
plt.legend(fontsize=14)
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

📍**Question 9:** 

A company is looking to build a new housing development in northern Arizona. The shallow subsurface in the area is composed of limestone and there has consequently been a history of large sinkholes. The company hires you to perform a gravity survey to look for shallow, underground caverns. 

The subsurface configurations in the figure below produce similar gravity anomaly curves. Which is more geologically plausible? Justify your answer with two geologic reasons at minimum! (2 points, 1 each)

**Note: The typical density of limestone is 2500-2800 $kg/m^3$.** And the **Density** column in the figure below marks the **density contrast**.

![](Question9.png)

_Type your answer here, replacing this text._

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are the lab 6 of ESS 314 Autumn 2023.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q2 (points: 1)
```
>>> np.issubdtype(type(contrast), np.number)
True
```
