```python
# Initialize Otter
import otter
grader = otter.Notebook("Lab7.ipynb")
```

# **ESS 314: LAB 7**
## Magnetics

**For this lab, you will need to submit this .ipynb file to the GradeScope as usual.** 

### ⚠️Before we get started, let's upgrade the Harmonica to have some of the new developed feature that is required for the lab. Run the cell below.

```python
!pip install git+https://github.com/fatiando/harmonica.git@c677dfa --upgrade
```

### ⚠️Make sure you are seeing a version: 'v0.6.0.post33+gc677dfa' below

```python
import numpy as np
import matplotlib.pyplot as plt
import verde as vd
import pandas as pd

import harmonica as hm
hm.__version__
```

---
Today we will look at some magnetic data. The magnetic method is quite complicated, so first we will do a few exercises to improve our understanding of the Earth’s magnetic field.

The intensity of the Earth’s magnetic field at some location on the Earth’s surface can be written as a vector $F_{E}$, defined by its inclination i (the angle the vector makes with a horizontal plane) and its declination d (the horizontal angle from geographic north). **Figure 1** below illustrates these concepts. If you find **Figure 1** a little confusing, then **Figure 2** below illustrates these concepts in a different way using a compass needle. Note that Z is positive down.

![](Fig1.png)
 <center> **Figure 1.** Diagram illustrating vector FE and its constituent vectors, along with declination and inclination angles d and i, respectively.

![](Fig2.png)
**Figure 2.** Representation of vector FE and its constituent vectors, along with declination and inclination angles d and i, respectively, using a compass schematic.

The total magnetic field can be broken down into its horizontal $H_E$ and vertical $Z_E$ components and be represented as

$$F_E = \sqrt{ (H_E)^2 + (Z_E)^2  }.$$ <div style="text-align: right"> (1) </div>
    
$H_E$ can be further resolved into northward and eastward components ($X_E$ and $Y_E$, respectively):

$$H_E = \sqrt{ (X_E)^2 + (Y_E)^2  }.$$ <div style="text-align: right"> (2) </div>
    
Combining equations (1) and (2), one can represent the magnetic field using
    
$$F_E = \sqrt{ (X_E)^2 + (Y_E)^2 + (Z_E)^2 }.$$ <div style="text-align: right"> (3) </div>


<!-- BEGIN QUESTION -->

🧭**Question 1:** 
Use simple trigonometry to derive the following formulas. We will do the expression for $H_E$ together in lab. Use the previous cell to learn how to write equations in Markdown. Markdown follows the LateX conventions. Enter an equation within the ``$ $`` to enter the mathematical environment, and follow LateX guidelines for the notations (operators, greek letters, powers and indices, fractions, roots, sums and integrals, brackets, ...) (https://en.wikibooks.org/wiki/LaTeX/Mathematics)

🧭a. Express $H_E$ and $Z_E$ in terms of the angle of inclination $i$ and the total field magnitude ($F_E$). (1 point)

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

🧭b. Express $X_E$ and $Y_E$ in terms of the total field magnitude ($F_E$) and the angles of inclination ($i$) and declination ($d$). (1 point)

_Type your answer here, replacing this text._

<!-- END QUESTION -->

🧭c. Assume the magnetic field vector is ($X_E$, $Y_E$, $Z_E$) = (4, 1, 1). Calculate the inclination ($i$) and declination ($d$) using the formulas you derived above. **Express your answer in degree.** (2 points)

```python
# Type your answer here, replacing "..." with your answer
inc = ... 
dec = ...
print("The inclination is %.4f degree." % inc)
print("The declination is %.4f degree." % dec)
```

```python
grader.check("q1c")
```

If the magnetic field of the earth could be produced by a simple north-south dipole (**Figure 3**), then inclination would be 0° anywhere along the equator, 90° (down) at the north pole and -90° (up) at the south pole.

![](Fig3.png)
**Figure 3.** The magnetic field around a sphere produced by a dipole.

However, the Earth’s magnetic field is imperfect. The magnetic equator (where i = 0°) is
irregular (**Figure 4**). Additionally, the inclination is ±90° at the north and south magnetic poles.
This is why **magnetic north does not equal true north.** 

![](Fig4.png)
**Figure 4.** Global magnetic inclination values. Source: https://www.ngdc.noaa.gov/geomag/WMM/

Therefore, while the Earth’s magnetic field is not perfect, it is close and is often approximated by a magnetic field produced by a dipole tilted 11.5° to the Earth’s axis of rotation. This reproduces 90% of the Earth’s magnetic field, and we can use this approximation for the vertical and horizontal components of the magnetic field at some point.

<!-- BEGIN QUESTION -->

🧭**Question 2:** In your own words, describe the distinctions between the geographic poles, the magnetic poles, and the geomagnetic poles. (2 points)

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---
A magnetic field is a potential field. The potential $V$ for a monopole is expressed by


\begin{align}
V = -\int_\infty^r \frac{m}{r^2} dr, \label{eq4}\tag{4} 
\end{align}
    
   
where $m$ is the pole strength and $r$ is the distance from the pole.


<!-- BEGIN QUESTION -->

🧭**Question 3:** Solve the integral to get a simpler equation for potential. Show your work.

Submit your math together with this notebook as Markdown/LaTeX, or as a separate document / image.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---
Without going through extensive derivations, the magnetic potential at point $P$ due to a dipole is

\begin{align} 
V = \frac{M\cos\theta}{r^2}, \label{eq5}\tag{4}
\end{align}


where $M$ is the magnetic moment, $\theta$ is the angle from the dipole to point $P$ (90-latitude for a geocentric dipole), and $r$ is the distance between point $P$ and the dipole. By taking the negative derivatives of the potential in spherical coordinates, one could derive the following equations for the radial and tangential components of the magnetic field:

\begin{align}
H_r &= -Z_E = \frac{2M\cos\theta}{r^3}. \label{eq6}\tag{6}\\
H_\theta &= -X_E = \frac{M\sin\theta}{r^3}.\label{eq7}\tag{7}
\end{align}

Because the Earth is spherical, the radial component $H_r$ is equivalent to the vertical field and the tangential component $H_\theta$ is equivalent to the horizontal southward pointing field. The E-W component of a dipole field ($Y_E$) is zero. See **Figure 5** below for a visualization of this.

![](Fig5.png)
**Figure 5.** Some important features of the Earth’s magnetic field.

🧭**Question 4a:** 

Assume a radius of $6371\;km$ for the Earth and a magnetic moment of the imaginary dipole responsible for the majority of Earth’s field as $8\times10^{25}$ emu.

🧭a. Estimate the horizontal $H_e$, vertical $Z_e$, and total $F_e$ fields at the 2006 North magnetic pole in nanoTeslas (latitude 80°; Figure 6). Note that $\theta$ is the colatitude (90 – latitude). The radius, $r$, must be expressed in centimeters. The answer you will get is expressed in oersteds. Multiply by $10^5$ to get the answer in nanoTeslas (nT).

```python
# Type your answer here, replacing "..." with your answer
Z_e = ...  # use equation (6)
H_e = ...  # use equation (7)
F_e = ...

print("The vertical component (Z_e) is %.3f nT" % Z_e)
print("The horizontal component (H_e) is %.3f nT" % H_e)
print("The total field (F_e) is %.3f nT" % F_e)
```

```python
grader.check("q4a")
```

<!-- BEGIN QUESTION -->

🧭**Question 4b:** 2 points

How well does it approximate the actual field strength as shown in Figure 6 below? What does the value represent?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

![](Fig6.png)
**Figure 6.** Values of total intensity for the geomagnetic field (FE) in 2006 for a portion of the Western Hemisphere.

<!-- BEGIN QUESTION -->

🧭**Question 5:** According to equation (6) and (7), what is the ratio of the dipole total field strength at the equator ($\theta=90^{\circ}$) to that at the pole ($\theta=0^\circ$)? Show your work. (2 points)

_Type your answer here, replacing this text._

<!-- END QUESTION -->

In a magnetic survey we look for perturbations between the observed magnetic field and a reference field. These perturbations are from smaller magnetic fields induced by geologic bodies.

The magnetic field that we actually measure is the sum of Earth’s magnetic field at the point of study and the anomalous field produced by a geologic body or bodies. However, we are interested in the total field anomaly in the direction of Earth’s main field. This is referred to as $F_{AT}$. $F_{AT}$ depends on the vertical and horizontal components of the anomalous field ($Z_A$ and $H_A$, respectively) and the inclination $i$:

\begin{align}
F_{AT} = Z_A\sin{i} + H_A\cos{i}.\label{eq8}\tag{8}
\end{align}

The next few questions are designed to help you gain some basic intuition on what the magnetic field induced in a simple geological body looks like.

<!-- BEGIN QUESTION -->

**Question 6:** At each location indicated below, determine whether the vertical or horizontal components of the anomalous magnetic field will be more important. Illustrate your answers using vectors.

a. Near the poles. *Hint: $i=90^\circ$*

b. Near the equator. *Hint: $i=0^\circ$*

_Type your answer here, replacing this text._

<!-- END QUESTION -->

You just proved that the component (vertical or horizontal) that affects the total induced field the most depends on latitude. Latitude also affects the orientation of Earth’s magnetic field, which induces the magnetic field in the geologic body. This gets confusing, because as you change latitude, not only are you seeing different contributions from the horizontal and vertical field in the total field, but you are also seeing changes in the horizontal and vertical fields themselves. **YIKES!**

---
## Buried Spherical Body

```python
# below we define a function that calculates the anomaly based on a buried magnetic sphere
# function takes the inputs z, R, k, x, inc
# function outputs the values of Z_a, H_a, and F_at

# z is the depth of the sphere in meters
# R is the radius of the sphere in meters
# k is the magnetic susceptibility in cgs emu (cm^(1/2)/s)
# x is the position on the surface in meters (x=0 is above center of sphere)
# inc is the inclination in degrees

# do not alter this part since this is where we define it
# skip ahead to the next cell to call the function

def mag_sphere(z,R,k,x,inc):
    # changes inclination if 0
    if inc == 0: # in DEGREES - changed to radians in script
        inc = 1e-10
    incr = np.deg2rad(inc) # convert degres to radian
    # calculate magnetic field strength
    theta = np.arctan(np.tan(incr)/2) #inclination to colatitude - this is in radians now
    eR = 6.4e8
    M = 8e25
    Z_e = ((2*M*np.cos(theta))/(eR**3))*(10**5)
    H_e = ((M*np.sin(theta))/(eR**3))*(10**5)
    F_e = np.sqrt(Z_e**2+H_e**2)
    
    Z_a = ((((4/3)*np.pi*R**3*k*F_e)*np.sin(incr))/(x**2+z**2)**(3/2))*(((3*z**2)/(x**2+z**2))-((3*x*z/np.tan(incr))/(x**2+z**2))-1);
    H_a = ((((4/3)*np.pi*R**3*k*F_e)*np.cos(incr))/(x**2+z**2)**(3/2))*(((3*z**2)/(x**2+z**2))-((3*x*z*np.tan(incr))/(x**2+z**2))-1);
    F_a = Z_a*np.sin(incr)+H_a*np.cos(incr)
    
    fig = plt.plot(x, Z_a, 'b:', label='Vertical')
    fig = plt.plot(x, H_a, 'r:', label='Horizontal')
    fig = plt.plot(x, F_a, 'g-', label='Total')
    plt.xlabel('Position (x)')
    plt.ylabel('Magnetic Field Anomaly (nT)')
    plt.legend()
    plt.show()
```

Now we will use the function mag_sphere to calculate, and plot the magnetif field anomaly components and total force.

```python
# input parameters
z = 5 # depth from the spherical anomaly in meters
R = 1 # radius of the spherical anomaly in meters
k = 0.005 # magnetic susceptibility in cgs emu (cm^(1/2)/s)
x = np.linspace(-50, 50, 101) # position on the surface in meters (x=0 is above center of sphere)
inc = 120 #degrees # inclination in degrees

# calling the function
mag_sphere(z, R, k, x, inc)
```

<!-- BEGIN QUESTION -->

🧭**Question 7:**  Use the `mag_sphere` function to show the horizontal, vertical, and total field curves for each of the following inclinations. For each plot, write a 1-2 sentence explanation describing the information shown in each plot and what it means for the magnetic anomalies you would see at those places on Earth's surface. **Only change the inclination parameter in the `mag_sphere` function and keep other paramemters unchanged.** (3 points)

a. north pole (inclination $i=90^\circ$)

b. Equator (inclination $i=0^\circ$)

c. Salem, OR (inclination $i=45^\circ$)

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

🧭**Question 8:** For the question below, **only change one parameter at a time in the `mag_sphere` function and keep other paramemters unchanged.** (2 points)

a. Move the spherical body deeper into the subsurface (increase $z$). What does this do to the curve?

b. Increase the susceptibility (increase $k$). What does this do to the total field?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

Despite the added complication of the latitude dependence of magnetic data, you can see that manipulating a simple rectangular body in the subsurface affects magnetic measurements similarly to how it affects gravity measurements. This also means that magnetic surveys have the same problem as gravity surveys: there are multiple subsurface models that can fit the same data set, the inverse problem is not unique. Therefore, the effectiveness of a magnetic study depends on whether there is previous geologic knowledge of the area.

---
## Real Buried Body: Magnetic Survey
Okay, lets finally get to an example. You are in Arizona (Latitude = 34°N) at an archeological site and are looking for old buried walls. Your colleagues already found one wall extending East/West, composed of basalt. Now you are conducting a magnetic survey to find more walls. You take 101 measurements in a South/North trending line. The results are in `ArchealogicalData.txt` that can be read through `pandas`. We will use Harmonica to model the magnetic data.

Similar to Harmonica for gravity modeling, we first need to define the coordinate system for observation (3-dimentional coordinate system). Scripts below define a `region` of 100x100 meters: 100 meters in both x and y direction (horizontal). The `shape` parameter defines who many points to measure in the defined region, which in this case, 101 measurement in both x and y direction. Then we would have gravity measurement at `x = [0,1,2,...,100]` and `y = [0,1,2,...,100]`.

```python
region = (0, 100, 0, 100) # (x_min, x_max, y_min, y_max), all in meters
shape = (101, 101)        # (n_x, n_y)
height = 0                # in meters
coordinates = vd.grid_coordinates(region, shape=shape, extra_coords=height)

prism = [0, 100, 44, 54, -10, -9] # (x_min, x_max, y_min, y_max, z_min, z_max), all in meters
magnetization = [0, 1, 0]        # vector in x, y, and z component, in epm (A/m)

# returned each component of the magnetic field generated by the blocks as arrays.
b_e, b_n, b_z = hm.prism_magnetic(coordinates, prism, magnetization)
```

We can plot the magnetic anomaly map (2D). We can also plot the anomaly profile (1D, along the red dashed line in the map).

```python
plt.figure(figsize=(14, 5), dpi=100)
plt.subplot(1,2,1)
plt.title("magnetic anomaly map (z direction)")
plt.pcolormesh(coordinates[0], coordinates[1], b_z)
plt.vlines(50, 1, 100, color='r', linestyle='--')
plt.colorbar(label="nT")
plt.xlabel("East/West (x, m)")
plt.ylabel("South/North (y, m)")

plt.subplot(1,2,2)
plt.title("magnetic anomaly profile")
plt.plot(b_z[:, 50], color='k', linestyle='-')
plt.xlabel("South/North (y, m)")
```

```python
# Read the data
df = pd.read_csv("./ArcheologicalData.txt", names=["measurement", "x", "b_z", "b_n", "b_e"])
df.head()        # show the first 5 rows
```

<!-- BEGIN QUESTION -->

🧭**Question 9:** 

Create a subsurface model that fits the magnetic data (you can only fit the z component). **Only change the shape of the magnetization block in y direction (north/south profile).** Two blocks can fit the data well. (2 points)

```python
# Type your answer here, replacing "..." with your answer
block = [[0, 100, ..., ..., ..., ...], # block 1: (x_min, x_max, y_min, y_max, z_min, z_max), all in meters
         [0, 100, ..., ..., ..., ...]] # block 2: (x_min, x_max, y_min, y_max, z_min, z_max), all in meters

magnetization = [[0, -1, 0],        # magnetization vector for block 1
                 [0, -1, 0]]        # magnetization vector for block 2
# returned each component of the magnetic field generated by the blocks as arrays.
b_e, b_n, b_z = hm.prism_magnetic(coordinates, block, magnetization)
```

```python
plt.figure(figsize=(10, 5), dpi=100)
plt.plot(df['x'], df['b_z'], label="field data")
plt.plot(df['x'], b_z[:, 50], label='my model')
plt.grid(True)
plt.legend(fontsize=12)
plt.ylabel("magnetic anomaly (nT)", fontsize=15)
plt.xlabel("South/North (y, m)", fontsize=15)
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

🧭**Question 10:**  Image that you performed both an induced magnetic survey and a gravity survey looking for features in the shallow subsurface. When you look at your data, the induced magnetism curve is completely flat, but the gravity survey reveals a large negative anomaly. What could this feature be? Why is there no induced magnetics? (1 point)

_Type your answer here, replacing this text._

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are the lab 7 of ESS 314 Autumn 2023.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1c (points: 2)
```
>>> np.issubdtype(type(inc), np.number)
True
```
```
>>> np.issubdtype(type(dec), np.number)
True
```


Tests for q4a (points: 3)
```
>>> np.issubdtype(type(Z_e), np.number)
True
```
```
>>> np.issubdtype(type(H_e), np.number)
True
```
```
>>> np.issubdtype(type(F_e), np.number)
True
```
